# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. Complete the code cells provided to you in this notebook, but do **not** change the names of the variables. If you do that, the autograder will fail and you will not receive any points.
2. Run all the answered cells before you run the testing cells. The answers must exist before they are graded!
3. Remove from each cell the code `raise NotImplementedError()` and replace it with your implementation.

## The dataset

The [`liste-des-immeubles-proteges-au-titre-des-monuments-historiques.csv`](data/liste-des-immeubles-proteges-au-titre-des-monuments-historiques.csv) file contains information about all the historical monuments (buildings) that can be found in the territory of France and are protected by the Government. These are the columns present in the dataset:

* `ref`: The building ID
* `tico`: The building's common name
* `adrs`: The location of the monument (the address)
* `affe`: Government agency in that is in charge of the monument
* `autr`: The name of the architect that built or designed the monument
* `com` and `wcom`: The [*commune*](https://en.wikipedia.org/wiki/Communes_of_France) where the building is located
* `desc`: Description
* `dpro`: Date of protection
* `dpt_lettre`: Province
* `hist`: History of the monument
* `insee`: [INSEE](https://www.insee.fr/en/accueil) code (ID given to the monument for statistical purposes)
* `ppro`: Legal details about the building's protection status
* `reg`: Region where the monument is located
* `sclx`: Historical period to which the monument belongs
* `stat`: An additional descriptor with information about the owner of the building

```python
import pandas as pd
monuments = pd.read_csv("data/liste-des-immeubles-proteges-au-titre-des-monuments-historiques.csv", sep=";")
monuments.head()
```

Output:
```
tico  \
0          Chapelle du cimetière   
1   Chapelle Notre-Dame-de-Santé   
2                         Maison   
3  Eglise paroissiale Saint-Loup   
4            Eglise Saint-Martin   

                                                ppro                     dpro  \
0  Chapelle du cimetière : classement par arrêté ...   1913/12/30 : classé MH   
1  Chapelle Notre-Dame-de-Santé : classement par ...   1932/12/29 : classé MH   
2  Ensemble de l'escalier avec rampe en fer forgé...  1948/04/13 : inscrit MH   
3  Eglise paroissiale : inscription par arrêté du...  1948/04/13 : inscrit MH   
4     Eglise : classement par arrêté du 13 juin 1913   1913/06/13 : classé MH   

                      stat desc  \
0  propriété de la commune  NaN   
1  propriété de la commune  NaN   
2         propriété privée  NaN   
3  propriété de la commune  NaN   
4  propriété de la commune  NaN   

                                                hist autr               adrs  \
0  Romane pour la plus grande par
```

### Exercise 1 (3 points)

Create a new column in the dataset that contains only the year in which the monument was protected. Call it `year_protection`.

**Hint:** If you choose to use the [`pd.to_datetime()`](https://pandas.pydata.org/docs/reference/api/pandas.to_datetime.html) function to convert from `str` to `datetime`, make sure to use the `errors="coerce"` argument, since a few rows have corrupted data (like references to the year 20115). This argument replaces those values with `NaT`s ("Not a Time"), which are the way `pandas` represents null values in Series of type `datetime`.

<!-- grade_id: cell-2402f86974fbcb94, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
monuments['dpro'] = pd.to_datetime(monuments['dpro'], errors='coerce')
monuments['year_protection'] = monuments['dpro'].dt.year
```

### Exercise 2 (2 points)

Using `.apply()` and/or ordinary string methods, create a new column called `castle` in the `monuments` DataFrame in which the values are `True` if the monument is a castle and `False` if it isn't, i.e., use the `tico` column and find whether the values contain the French word for castle. "Castle" in French is *château*.

**Hint:** There are 6,034 castles in this dataset, and 39,650 buildings that are not castles. To compute these numbers, we accounted for variation due to uppercase and lowercase differences and the common typo *chateau* (without the circumflex). Your results should correspond with this.

<!-- grade_id: cell-0b595804775e26d4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
monuments['castle']=monuments['tico'].apply(lambda x:'château' in x.lower().replace('chateau','château'))
```

### Exercise 3 (1 point)

The `tico` column has many duplicates. For instance, *Maison*, which means "House" in French, is repeated many times because there are many monuments that are simply known as "houses". Create a **list** (*not* a new column in the `monuments` DataFrame) called `tico_once`, which exclusively lists the elements of the `tico` column that appear only *once*. This means that a mere `.drop_duplicates()` won't be valid here, because it would still include words such as "maison" or "château", that show up many times.

**Hint:** You should: Make all words lowercase, count how many times each row appears, remove those that appear more than once, and assign the remaining to the `tico_once` list. [Here](https://stackoverflow.com/a/28272238/6945498) is how to filter a Series in case you need to.

<!-- grade_id: cell-36e3ac92b4e3f7f3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
monuments['tico'] = monuments['tico'].str.lower()
counts = monuments['tico'].value_counts()
tico_once = counts[counts==1].index.tolist()
```

## Testing cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get an error when checking your answers.

<!-- grade_id: cell-77d5c96556a2fd02, points: 3, grade: true, solution: false -->
```python
# Ex 1
assert "year_protection" in monuments.columns, "Ex. 1 - You have to create a column called year_protection in the monuments DataFrame!"
print("Exercise 1 passed the preliminary sanity check.")
```

Output:
```
Exercise 1 passed the preliminary sanity check.
```

<!-- grade_id: cell-eba8ec71dfdc1b4d, points: 2, grade: true, solution: false -->
```python
# Ex 2
assert "castle" in monuments.columns, "Ex. 2 - You have to create a column called castle in the monuments DataFrame!"
print("Exercise 2 passed the preliminary sanity check.")
```

Output:
```
Exercise 2 passed the preliminary sanity check.
```

<!-- grade_id: cell-91a308715ed234fb, points: 1, grade: true, solution: false -->
```python
# Ex 3
assert "tico_once" in globals(), "Ex. 3 - You need to create the tico_once variable!"
assert type(tico_once) == type([1,2,3]), "Ex. 3 - The tico_once variable must be of type list!"
assert len(tico_once) == 23874, "Ex. 3 - The tico_once list must have exactly 23874 elements. Are you standardizing the case? No need to correct typos or replace special characters."
print("Exercise 3 passed the preliminary sanity check.")
```

Output:
```
Exercise 3 passed the preliminary sanity check.
```

## Attribution

"Immeubles protégés au titre des Monuments Historiques", July 23, 2020, 
Ministère de la Culture, [ETALAB 2.0 license](https://www.etalab.gouv.fr/wp-content/uploads/2017/04/ETALAB-Licence-Ouverte-v2.0.pdf), https://www.data.gouv.fr/fr/datasets/immeubles-proteges-au-titre-des-monuments-historiques-2/